Bevor Sie dieses Notebook abgeben, testen Sie, ob alles so funktioniert, wie Sie es erwarten. So sollten Sie z.B. den ** Kernel neu starten** (im Menu w&auml;hlen Sie Kernel$\rightarrow$Restart) und dann ** alle Zellen ausf&uuml;hren **  (im Menu w&auml;hlen Sie Cell$\rightarrow$Run All).

F&uuml;llen Sie alle Stellen, die mit  `YOUR CODE HERE` or "YOUR ANSWER HERE" aus. Sie d&uuml;rfen zus&auml;tzlich eigene Zellen und Funktionen definieren, nicht jedoch die Signaturen der gegebenen Funktionsr&uuml;mpfe &auml;ndern.

Bitte beachten Sie auch die "Hinweise zu den Abgaben"-Datei in Moodle.

```python
NAME = ""
```

---

<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"></ul></div>

```python
import numpy as np
import itlinsta18 as itlin
import scipy.linalg as spla
```

Ein Euler-Bernpoulli-Balken ist ein einfaches Modell für die Biegung eines Balkens der Länge $L$ bei Belastung. 

Die vertikale Auslenkung sei gegeben durch die Funktion $y(x)$, $0\leq x\leq L$.
$$ EIy''''(x)=f(x) $$
wobei $E$ das Elastizitätsmodul (Materialkonstante) und $I$ das Flächenträgheitsmoment des Querschnitts ist 
(eine geometrische Größe). Beide sind konstant entlang des Balkens. 

$f(x)$ ist die Kraft pro Längeneinheit, die sich entlang des Balkens ändern kann.


Je nach Art der Einspannung/ des Auflagers erhält man verschiedene 
 Randwertprobleme (s.u.). Diese werden meist durch Finite Differenzen-Verfahren (s. letztes Kapitel) gelöst.

Zur numerischen Behandlung  wird die 4.Ableitung   durch die Differenzenapproximation
$$ y''''(x)= \frac{y(x-2h)-4y(x-h)+6y(x)-4y(x+h)+y(x+2h)}{h^4} $$
angenähert und in der Dgl. ersetzt. 
Dazu wird der Balken in $n+1$ Segmente unterteilt.
Es sei $h=\frac L{n+1}$ die Segmentbreite.
Für die Diskterisierungspunkte erhält man
$0=x_0<x_1<\dots <x_n<x_{n+1}=L$ mit $h=x_i-x_{i-1}$. 
Ersetzt man nun  $y''''$ durch die Differenzenapproximation, so erhält man
$$ y_{i-2}-4y_{i-1}+6y_i-4y_{i+1}+y_{i+2} = \frac{h^4}{EI}f(x_i), \qquad i=2,\dots,n-1 $$

Wir betrachten einen Balken, der an beiden Enden eingespannt ist: $$ y(0)=y'(0)=y(L)=y'(L)=0$$
Für die Randbedingungen erhält man mit Hilfe der Annäherung:
$$ y''''(x)\approx \frac{12y(x+h)-6y(x+2h)+\frac 4 3 y(x+3h)}{h^4} $$
für $x=0$. Diese Näherung gilt, falls $y(x)=y'(x)=0$ 

Insgesamt erhält man damit als Gleichungssystem
$$ \begin{pmatrix}
12&-6&\frac 4 3 \\
-4&6&-4&1\\
1&-4&6&-4&1\\
&& \ddots &\ddots &\ddots &\ddots &\ddots\\
&&&1&-4&6&-4&1\\
&&&&1&-4&6&-4&1\\
&&&&&1&-4&6&-4\\
&&&&&&\frac 4 3&-6&12
\end{pmatrix}
\begin{pmatrix}y_1\\ y_2\\ \vdots\\ \\ \vdots\\ y_{n-1}\\ y_n
\end{pmatrix}
=\frac{h^4}{EI}
\begin{pmatrix}f(x_1)\\ f(x_2)\\ \vdots \\ \\ \vdots\\f(x_{n-1})\\ f(x_n)
\end{pmatrix}
$$
Die Werte an den Rändern ergeben sich aus den Randbedingungen.

Geben Sie eine Begründung für die zweite und vorletzte Zeile des Gleichungssystems.

<!-- grade_id: cell-0b7ba6a1544d3c4a, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Welche Konvergenzaussagen für Jacobi- und Gauss-Seidel-Verfahren für dieses Problem können Sie vorab treffen?

<!-- grade_id: cell-632423ee9e87d675, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Wir betrachten als erstes einen Stahlbalken der Länge $L=15 m$, der  Tiefe $d=7cm$ und Breite $b=10cm$. The Dichte von Stahl ist
ungefähr $\rho=7850 kg/m^3$, $E=2\cdot 10^{11}$ N/m$^2$. $I=\frac{bd^3}{12}$.

Zuerst nehmen wir zusätzlich an, dass kein zusätzliches Gewicht den Balken belastet, so dass $f(x)$ nur das Eigengewicht des Balkens repräsentiert. 

Das Gewicht $f$ des Balkens pro $m$ ist somit $\rho\cdot b \cdot d$. Die Gewichtskraft ergibt sich daraus durch Multiplikation mit der Graviatationsbeschleunigung $g= -9.81 m/s^2$.

Erstellen Sie eine Funktion `asparse(n)`, die eine dünnbesetzte Version der Matrix (coo_matrix) in Abhängigkeit von $n$ berechnet und zurückgibt. Die Matrix hat dann die Dimension $n\times n$

<!-- grade_id: cell-0e10af5b2ed00f8b, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0f857e1f1c15580e, points: 1, grade: true, solution: false -->
```python
# Autograder Tests
assert np.allclose(asparse(8).todense(), itlin.asparse(8).todense())
```

Implementieren Sie je eine Routine `jacobi_beam(y, c)` bzw. `gs_beam(y, c)`, die  einen Iterationsschritt für das Jacobi- bzw. das Gauss-Seidel-Verfahren implementiert. Implementieren Sie die Verfahren Rechenzeit- und speicherplatzsparend. 
Dies bedeutet insbesondere, dass Sie nirgendwo die komplette Matrix (auch nicht in dünnbetzter Form) speichern, sondern die Verfahren direkt so implementieren, dass es auf genau diese Struktur 
der Matrix passt.

<!-- grade_id: cell-de740f04bf00c4ee, points: 4, grade: true, solution: true -->
```python
def jacobi_beam(y, c):
    """
    Jacobi adapted for Euler-Bernoulli-beam
    One iteration only
    y: approximation to solution
    c: right hand side
    return: next iterate
    """
    # YOUR CODE HERE
    raise NotImplementedError()


def gs_beam(y, c):
    """
    Gauss-Seidel adapted for Euler-Bernoulli-beam
    One iteration only
    y: approximation to solution
    c: right hand side
    return: next iterate
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-18e017d87d83b754, points: 1, grade: true, solution: false -->
```python
# Autograder tests - Zelle nicht verändern oder löschen
y = np.ones(6)
c = np.ones(6)
assert np.allclose(jacobi_beam(y, c), itlin.jacobi_beam(y, c))
```

<!-- grade_id: cell-e9689e570ead3891, points: 1, grade: true, solution: false -->
```python
# Autograder tests - Zelle nicht verändern oder löschen
y = np.ones(6)
c = np.ones(6)
assert np.allclose(gs_beam(y, c), itlin.gs_beam(y, c))
```

<!-- grade_id: cell-96fa4b6e519a40ba, points: 1, grade: true, solution: false -->
```python
# Hidden tests - Zelle nicht verändern oder löschen
# BEGIN  HIDDEN TESTS
y = np.random.rand(7)
c = np.random.rand(7)
assert np.allclose(jacobi_beam(y, c), itlin.jacobi_beam(y, c))
```

<!-- grade_id: cell-fb17a2c26dc65552, points: 1, grade: true, solution: false -->
```python
# Hidden tests - Zelle nicht verändern oder löschen
# BEGIN  HIDDEN TESTS
y = np.random.rand(7)
c = np.random.rand(7)
assert np.allclose(gs_beam(y, c), itlin.gs_beam(y, c))
```

Schreiben Sie eine Routine `multi_it`, die die Iterationsschleife implementiert und die die oben von Ihnen implementierten Funktionen in jedem Iterationsschritt aufruft.
Die Iteration soll abgebrochen werden, wenn die maximale Iterationszahl überschritten wird oder die Norm des Inkrements kleiner als die geg. Toleranz ist. 


<!-- grade_id: cell-be61852548315fd7, points: 2, grade: true, solution: true -->
```python
def multi_it(method, n, c, tol=1e-5, itmax=1000):
    """ 
    iteration loop over method (gauss seidel or jacobi) 
    method: gs or jacobi
    n: system dimension
    c: right hand side
    tol: tolerance (increment)
    itmax: maximum number of iterations
    return: 
        y: last iterate
        it: number of iterations used so far
        enorm: norm of increment
        kfac: Konvergenzrate
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Implementieren Sie die Funktion $f$.

<!-- grade_id: cell-ffbea65962b1e766, points: 1, grade: true, solution: true -->
```python
# import packages
import numpy as np

# define global constants b, d, L, E, I, g, rho load
d = 0.05
b = 0.1
L = 15
E = 2E11
I = b*d**3/12
g = 9.81
rho = 7850


def f(x):
    """
    load function, right hand side f
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f456cca875a2127c, points: 1, grade: true, solution: false -->
```python
# Autograder Tests
assert np.allclose(f(5.3), itlin.f(5.3))
```

Lösen Sie das System für $n=15$ mit den beiden Verfahren. 
Konvergieren die Verfahren? Falls ja, wie viele Iterationen sind jeweils notwendig, damit die Abbruchgenauigkeit mindestens $10^{-5}$ ist?

<!-- grade_id: cell-e9dc32b80c69289a, points: 1, grade: true, solution: true -->
```python
# Ihr Code für den Aufruf der  Jacobi-Iteration für das Testbeispiel mit n=15
# YOUR CODE HERE
raise NotImplementedError()
```

Ihr Kommentar zu Ergebnissen/Verhalten der Jacobi-Iteration

<!-- grade_id: cell-465dd5108cc133ca, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-32347a4bf4be4e09, points: 2, grade: true, solution: false -->
```python
# Hidden tests Jacobi
```

<!-- grade_id: cell-6a2d6184b92ee8ac, points: 1, grade: true, solution: true -->
```python
# Ihr Code für den Aufruf der Gauss-Seidel-Iteration für das Testbeispiel mit n=15
# YOUR CODE HERE
raise NotImplementedError()
```

Ihr Kommentar zum Gauss-Seidel-Ergebnis/Verhalten. (Plots dazu kommen weiter unten)

<!-- grade_id: cell-4824d86cd62d0467, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-84a5a4ece92cb7fb, points: 3, grade: true, solution: false -->
```python
# Hidden Tests Gauss-Seidel - Zelle nicht ändern
```

Plotten Sie die Lösung gegen die exakte Lösung $y(x)=\frac{f(x)x^2(L-x)^2}{24EI}$. Wie viele Iterationen und welche Toleranz sind notwendig, damit die Lösungen im Intervallmittelpunkt weniger als $10^{-2}$ auseinanderliegen?

<!-- grade_id: cell-209b7955c0bf9f12, points: 3, grade: true, solution: true -->
```python
# Plot
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c716b420dc01e603, points: 3, grade: true, solution: true -->
```python
# Bestimmung der Anzahl der Iterationen
# YOUR CODE HERE
raise NotImplementedError()
```

Nun wird der Balken durch ein Gewicht von $2000 kg$ belastet, das gleichmäßig verteilt zwischen $x=8$ und $x=10$ auf dem Balken liegt, so dass Sie in diesem Bereich $f$ entsprechend modifizieren müssen.

Lösen Sie das Problem wieder mit demselben  $n$. Plotten Sie die Lösung. Wo hat der Balken die größte Auslenkung?

<!-- grade_id: cell-5f353f24955a3bf0, points: 2, grade: true, solution: true -->
```python
# Balken mit Zusatzgewicht
# YOUR CODE HERE
raise NotImplementedError()
```

Nun soll das ursprüngliche System mit einem cg-Verfahren gelöst werden.
Ws können Sie vorab über die Konvergenz sagen?

<!-- grade_id: cell-3cb25835a1089d1c, points: 1, grade: true, solution: true -->
```python
# Konvergenzaussagen mathematisch prüfen.
# Textantwort in die nächste Zelle.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b6aea173382fa06a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Nun soll das ursprüngliche System mit dem cg-Verfahren gelöst werden.
Schreiben Sie dazu eine Funktion `amaly(y)`, die die Matrix-Vektor-Multiplikation $A\cdot y$ für den Euler-Bernopulli-Balken effizient implementiert. Rückgabewert ist der Vektor $A\cdot y$.

Vergessen Sie nicht, das cg-Verfahren nach $n$ Iterationen neu zu starten. 

Eine effiziente Impementation ist wichtig!

Vergleichen Sie das Verfahren bzgl. Schrittzahl und Rechenzeit mit dem Gauss-Seidel-Verfahren. 

<!-- grade_id: cell-f0d07d165190d0aa, points: 3, grade: true, solution: true -->
```python
import scipy.linalg as spla
import scipy as sp


def cg(amaly, c, y, itmax=100, eps=1e-5):
    """
    in: 
    amaly: Funktion, die A*y effizient für dieses konkrete Problem berechnet
    c: rechte Seite
    y: Startlösung
    itmax: maximal zulässige Anzahl Iterationen
    eps: Toleranz für die Norm des Residuums
    return:
    yn: Lösung bzw. aktuelle Iterierte bei Nichtkonvergenz
    it: Anzahl verwendeter Iterationen
    r: Residuum
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return xn, it, r


def amaly(y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4b2bcb29a68caa47, points: 1, grade: true, solution: false -->
```python
# Hidden test - nicht löschen
```

<!-- grade_id: cell-e83f63e64265bcd3, points: 1, grade: true, solution: false -->
```python
# Hidden test - nicht löschen
```

<!-- grade_id: cell-16d02584539938bf, points: 2, grade: true, solution: false -->
```python
# Hidden test - nicht löschen oder verändern
n = 15
```

<!-- grade_id: cell-ed21ee6e2a38f509, points: 2, grade: true, solution: false -->
```python
# Hidden test - nicht löschen
```

Meine Kommentare zur Implementation und zum Aufwandsvergleich:

<!-- grade_id: cell-d4cfee1d8b3a32a1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-b0519811166eab19, points: 2, grade: true, solution: false -->
```python
# Hidden test - nicht löschen
```